---
id: 66b102ee0bdbad34a9f85ba0
title: Step 1
challengeType: 1
dashedName: step-1
lang: pt-BR
translated: true
updatedAt: '2025-09-29T05:49:28.202Z'
---

# --description--

In this workshop, you are going to continue learning about strings by building a Teacher Chatbot. 

To begin, add a `console` statement, with the message of `"Hi there!"`.

# --hints--

You should have a `console` statement.

```js
assert.match(code, /console\.log(.*)/);
```

Your console statement should output the message `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
